<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math"><article class="exercise exercise-like"><div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="exercise-python-including">
<div class="ac_question col-md-12" id="exercise-python-including_question"><div class="para" id="p-22">Compute the total amount of money loaned and store it in the variable <code class="code-inline tex2jax_ignore">loan_total</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="exercise-python-including_editor" data-question_label="" data-codelens="true" data-include="statistics">loan_total = 0
for loan in loan_amount:
    loan_total += loan
print(loan_total)
====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):

    def testOne(self):
        self.assertTrue('loan_total' in self.getEditorText(), "you need a loan_total variable")
        self.assertEqual(loan_total, sum(loan_amount), "Use the accumulator pattern to add up all the loans")
        self.assertFalse('sum(' in self.getEditorText(), "you may not use sum()")

MyTests().main()
</textarea>
</div></div></div></article></body>
</html>
